# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD25"
COLLABORATORS = "Adam Lewiński 407657, Małgorzata Serwańska 405044"
```

---

## Zestaw zadań 1: Elementy statystyki opisowej 1 (Notatnik 2/4)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import dataframe_image as dfi
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Wczytanie danych
Wczytaj dane z przygotowanego wcześniej pliku eruptions_data.csv do tablicy `eruptions_data` w taki sposób, aby kolumna Eruption Number stanowiła indeks tablicy.

Wyświetl przygotowaną tablicę.

<!-- grade_id: cell-184ec13f50fb3bfe, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
eruptions_data = pd.read_csv("eruption_data.csv", index_col = "Eruption Number")
print(eruptions_data)
```

Output:
```
Volcano Name Eruption Category  VEI  Start Year  \
Eruption Number                                                           
22486                       Cotopaxi         Confirmed    2        2022   
22481                           Taal         Confirmed    1        2022   
22458                      Turrialba         Confirmed    1        2022   
22453                         Ulawun         Confirmed    2        2022   
22454                          Raung         Confirmed    2        2022   
...                              ...               ...  ...         ...   
14556                      Tongariro         Confirmed    5       -9650   
21101            Craters of the Moon         Confirmed    0      -10060   
22141                   Igwisi Hills         Confirmed    1      -10450   
22351                   Quetrupillan         Confirmed    3      -10658   
22352                   Quetrupillan         Confirmed    3      -11345   

                 Start Month  Start Day  End Yea
```

### Zadanie 2: Przegląd danych - wybrane elementy [5 pkt]

#### a) Największe erupcje
Przygotuj i wyświetl tablicę `big_eruptions_data` zawierającą tylko i wyłącznie wiersze tablicy `eruptions_data` dotyczące erupcji, którym została przyporządkowana wartość VEI równa 7.

<!-- grade_id: cell-ae4385a252fa931d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
big_eruptions_data = eruptions_data[eruptions_data["VEI"] == 7]
df_styled = big_eruptions_data.style.format(precision=0)
dfi.export(df_styled, 'najwiekszeErupcje.png')
display(df_styled)
```

Output:
```
error: unknown flag `disable-gpu'
File ~/.local/lib/python3.10/site-packages/dataframe_image/_screenshot.py:142, in Screenshot.generate_image_from_html(self, args)
    141 def generate_image_from_html(self, args):
--> 142     subprocess.run(executable=self.chrome_path, args=args, check=True)

File /usr/lib/python3.10/subprocess.py:524, in run(input, capture_output, timeout, check, *popenargs, **kwargs)
    522     retcode = process.poll()
    523     if check and retcode:
--> 524         raise CalledProcessError(retcode, process.args,
    525                                  output=stdout, stderr=stderr)
    526 return CompletedProcess(process.args, retcode, stdout, stderr)

CalledProcessError: Command '['--enable-logging', '--disable-gpu', '--headless', '--no-sandbox', '--crash-dumps-dir=/tmp', '--force-device-scale-factor=1', '--window-size=1400,900', '--hide-scrollbars', '--screenshot=/tmp/tmps1w4vs0g/temp.png', '/tmp/tmps1w4vs0g/temp.html']' returned non-zero exit status 64.
```

<!-- grade_id: cell-93d159a92fa925b4, points: 1, grade: true, solution: false -->
```python
### Komórka testowa
assert set(big_eruptions_data['Volcano Name'].index) == set([16231, 20843, 13879, 20904, 16980, 20610, 18903])
assert set(big_eruptions_data['Volcano Name'].values) == set(['Tambora', 'Rinjani', 'Santorini', 'Blanco, Cerro', 'Kikai', 'Crater Lake', 'Kurile Lake'])
```

#### b) Najaktywniejsze wulkany
Przygotuj tablicę `most_active_data` zawierającą następujące elementy:

Indeks:
- domyślny

Kolumny:
- Volcano - nazwy wulkanów
- Recorded eruptions - liczba zarejestrowanych erupcji

Dane w tablicy powinny być posortowane od największej do najmniejszej liczby zarejestrowanych erupcji, a w przypadku takiej samej liczby erupcji alfabetycznie po nazwach wulkanów.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było 10 pierwszych wierszy, a zawarte w tablicy liczby wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-6822b057abf5c287, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
most_active_data = pd.DataFrame()
most_active_data["Volcano"] = eruptions_data["Volcano Name"]
most_active_data = most_active_data.groupby(by="Volcano", as_index=False).size()
most_active_data.columns = ["Volcano", "Recorded eruptions"]
most_active_data = most_active_data.sort_values(["Recorded eruptions", 'Volcano'], ascending = [False, True])
most_active_data.reset_index(inplace = True)
df_styled = most_active_data.head(10).style.format(precision=0)
dfi.export(df_styled, 'mostActiveData.png')
display(df_styled)
```

<!-- grade_id: cell-7de8c79d77dbc9db, points: 2, grade: true, solution: false -->
```python
### Komórka testowa
assert list(most_active_data.loc[:9, "Volcano"]) == ['Fournaise, Piton de la', 'Asosan', 'Villarrica', 'Asamayama', 'Etna',
                                                         'Mauna Loa', 'Klyuchevskoy','Sheveluch', 'Gamalama', 'Merapi']
assert list(most_active_data.loc[:9, "Recorded eruptions"]) == [191, 180, 152, 128, 121, 110, 104, 102,  81,  81]
```

#### c) Ostatnie erupcje o określonej wielkości
Przygotuj zmienną `last_eruption_data` zawierającą tablicę zawierającą następujące elementy:

Indeks:
- VEI - wartości VEI posortowane od najmniejszej do największej.

Kolumny:
- Last recorded eruption - lata, w których ostatni raz zarejestrowano erupcję o określonej wartości VEI (na podstawie daty początku erupcji).

Wyświetl przygotowaną tablicę w taki sposób aby zawarte w niej liczby wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-f26a7582febc53a6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
last_eruption_data = eruptions_data[["VEI", "Start Year"]]
last_eruption_data = last_eruption_data.groupby("VEI").max()
last_eruption_data.columns = ["Last recorded eruption"]
df_styled = last_eruption_data.style.format(precision=0)
dfi.export(df_styled, 'lastErutpionData.png')
display(df_styled)
```

<!-- grade_id: cell-d6641c86a43e0ab5, points: 2, grade: true, solution: false -->
```python
### Komórka testowa
assert (last_eruption_data.index.name == 'VEI') == True
assert list(last_eruption_data['Last recorded eruption'].index) == [0, 1, 2, 3, 4, 5, 6, 7]
assert list(last_eruption_data['Last recorded eruption'].values) == [2021, 2022, 2022, 2021, 2021, 2021, 1991, 1812]
```

### Zadanie 3: Konstrukcja szeregów rodzielczych - wersja podstawowa [10 pkt]

#### a) Forma tabelaryczna
Przygotuj szeregi rozdzielcze VEI w formie tablic zawierających następujące elementy:

Indeks:
- VEI - wartości VEI posortowane od najmniejszej do największej.

Kolumny:
- f - zliczenia poszczególnych kategorii (ang. frequency, absolute frequency),
- cf - zliczenia skumulowane poszczególnych kategorii (ang. cumulative frequency, cumulative absolute frequency),
- rf - częstość zliczeń poszczególnych kategorii (ang. relative frequency),
- crf - skumulowana częstość zliczeń poszczególnych kategorii (ang. cumulative relative frequency).

Przygotuj osobne tablice dla wszystkich erupcji (`vei_data_1`), erupcji potwierdzonych (`vei_data_2`) i erupcji niepotwierdzonych (`vei_data_3`).

Wyświetl przygotowane tablicę w taki sposób, aby liczby w kolumnach f i cf oraz indeks tablicy VEI wyświetlały się bez miejsc po przecinku, a liczby z kolumn rf i crf z dokładnością do 4 miejsca po przecinku.

<!-- grade_id: cell-10bd4bdf7a80603c, points: 0, grade: false, solution: true -->
```python
# Wszystkie erupcje
# YOUR CODE HERE
vei_data_1 = eruptions_data[["VEI", "Eruption Category"]]
vei_data_1 = vei_data_1.groupby("VEI").count()
cf = []
for index, row in vei_data_1.iterrows():
    try:
        cf.append(cf[-1]+row["Eruption Category"])
    except:
        cf.append(row["Eruption Category"])
vei_data_1["2"] = cf
vei_data_1["3"] = vei_data_1["Eruption Category"]/np.sum(vei_data_1["Eruption Category"])
vei_data_1["4"] = vei_data_1["2"]/np.sum(vei_data_1["Eruption Category"])
vei_data_1.columns = ["f","cf","rf","crf"]
vei_data_1.style.format(precision=0).format(precision=4, subset=["rf","crf"])
```

<!-- grade_id: cell-bbe66d65d1b300ca, points: 0, grade: false, solution: true -->
```python
# Potwierdzone erupcje
# YOUR CODE HERE
vei_data_2 = eruptions_data[eruptions_data["Eruption Category"] == "Confirmed"][["VEI", "Eruption Category"]]
vei_data_2 = vei_data_2.groupby("VEI").count()
cf = []
for index, row in vei_data_2.iterrows():
    try:
        cf.append(cf[-1]+row["Eruption Category"])
    except:
        cf.append(row["Eruption Category"])
vei_data_2["2"] = cf
vei_data_2["3"] = vei_data_2["Eruption Category"]/np.sum(vei_data_2["Eruption Category"])
vei_data_2["4"] = vei_data_2["2"]/np.sum(vei_data_2["Eruption Category"])
vei_data_2.columns = ["f","cf","rf","crf"]
vei_data_2.style.format(precision=0).format(precision=4, subset=["rf","crf"])
```

<!-- grade_id: cell-0d54ec5c8d49d655, points: 0, grade: false, solution: true -->
```python
# Niepotwierdzone erupcje
# YOUR CODE HERE
vei_data_3 = eruptions_data[eruptions_data["Eruption Category"] == "Uncertain"][["VEI", "Eruption Category"]]
vei_data_3 = vei_data_3.groupby("VEI").count()
for i in range(len(vei_data_1)-len(vei_data_3)):
    vei_data_3 = pd.concat([vei_data_3, pd.DataFrame({"Eruption Category": 0}, index = [vei_data_3.index[-1]+1,])])
vei_data_3.index.name="VEI"
cf = []
for index, row in vei_data_3.iterrows():
    try:
        cf.append(cf[-1]+row["Eruption Category"])
    except:
        cf.append(row["Eruption Category"])
vei_data_3["2"] = cf
vei_data_3["3"] = vei_data_3["Eruption Category"]/np.sum(vei_data_3["Eruption Category"])
vei_data_3["4"] = vei_data_3["2"]/np.sum(vei_data_3["Eruption Category"])
vei_data_3.columns = ["f","cf","rf","crf"]
vei_data_3.style.format(precision=0).format(precision=4, subset=["rf","crf"])
```

<!-- grade_id: cell-37f309620c903da9, points: 5, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1: Wszystkie erupcje
assert (vei_data_1.index.name == 'VEI') == True
assert len(set(list(vei_data_1.columns)) & set(['f', 'cf', 'rf', 'crf'])) == 4
assert list(vei_data_1['f'].index) == [0, 1, 2, 3, 4, 5, 6, 7]
assert list(vei_data_1['f'].values) == [1007, 1384, 3991, 1152, 511, 180, 53, 7]
# Test 2: Tylko potwierdzone erupcje
assert (vei_data_2.index.name == 'VEI') == True
assert len(set(list(vei_data_2.columns)) & set(['f', 'cf', 'rf', 'crf'])) == 4
assert list(vei_data_2['f'].index) == [0, 1, 2, 3, 4, 5, 6, 7]
assert list(vei_data_2['f'].values) == [803, 1227, 3727, 1133, 510, 180, 53, 7]
# Test 3: Tylko niepotwierdzone erupcje
assert (vei_data_3.index.name == 'VEI') == True
assert len(set(list(vei_data_3.columns)) & set(['f', 'cf', 'rf', 'crf'])) == 4
assert list(vei_data_3['f'].index) == [0, 1, 2, 3, 4]
assert list(vei_data_3['f'].values) == [204, 157, 264, 19, 1]
```

#### b) Forma graficzna
Zwizualizuj dane z przygotowanych tablic za pomocą wykresów słupkowych zliczeń (lewa oś wykresu) z nałożonymi na nie łamanymi częstości skumulowanej (prawa oś wykresu). Zestaw wykresy na pojedynczej figurze (siatka wykresów 1x3). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo ustaw:
- identyczne zakresy wartości na odpowiadających sobie osiach wszystkich trzech wykresów,
- różne od domyślnych i od siebie nawzajem kolory wykresu słupkowego oraz łamanej częstości,
- kolory lewej i prawej osi oraz ich opisów identyczne z kolorem elementu wykresu, którego dotyczą,
- format wyświetlania wartości VEI bez miejsc po przecinku.

<!-- grade_id: cell-4291e1fff2828e80, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
fig, ax = plt.subplots(1,3,figsize=(14,5), facecolor="white")
ax[0].bar(vei_data_1.index, vei_data_1.f, color = "green")
ax[0].set_ylabel("Frequency")
ax[0].set_ylim([0,4200])
ax[0].yaxis.label.set_color('green')
ax[0].tick_params(axis='y', colors='green')
ax[0].set_xlabel("VEI")
ax[0].set_xticks(range(0,8))
ax[0].set_facecolor("white")
ax[0] = ax[0].twinx()
ax[0].set_ylabel("Cumulated relative frequency")
ax[0].set_ylim([0,1.1])
ax[0].yaxis.label.set_color('red')
ax[0].tick_params(axis='y', colors='red')
ax[0].plot(vei_data_1.index, vei_data_1.crf, c="red")
ax[0].spines['left'].set_color('green')
ax[0].spines['right'].set_color('red')
ax[0].set_title("Wszystkie erupcje")
ax[0].set_xlim([-0.5,7.5])


ax[1].bar(vei_data_2.index, vei_data_2.f, color = "green")
ax[1].set_ylabel("Frequency")
ax[1].set_ylim([0,4200])
ax[1].yaxis.label.set_color('green')
ax[1].tick_params(axis='y', colors='green')
ax[1].set_xlabel("VEI")
ax[1].set_xticks(range(0,8))
ax[1].set_facecolor("white")
ax[1] = ax[1].twinx()
ax[1].set_ylabel("Cumulated relative frequency")
ax[1].set_ylim([0,1.1])
ax[1].yaxis.label.set_color('red')
ax[1].tick_params(axis='y', colors='red')
ax[1].plot(vei_data_2.index, vei_data_2.crf, c="red")
ax[1].spines['left'].set_color('green')
ax[1].spines['right'].set_color('red')
ax[1].set_title("Potwierdzone erupcje")
ax[1].set_xlim([-0.5,7.5])


ax[2].bar(vei_data_3.index, vei_data_3.f, color = "green")
ax[2].set_ylabel("Frequency")
ax[2].set_ylim([0,4200])
ax[2].yaxis.label.set_color('green')
ax[2].tick_params(axis='y', colors='green')
ax[2].set_xlabel("VEI")
ax[2].set_xticks(range(0,8))
ax[2].set_facecolor("white")
ax[2] = ax[2].twinx()
ax[2].set_ylabel("Cumulated relative frequency")
ax[2].set_ylim([0,1.1])
ax[2].yaxis.label.set_color('red')
ax[2].tick_params(axis='y', colors='red')
ax[2].plot(vei_data_3.index, vei_data_3.crf, c="red")
ax[2].spines['left'].set_color('green')
ax[2].spines['right'].set_color('red')
ax[2].set_title("Niepotwierdzone erupcje")
ax[2].set_xlim([-0.5,7.5])

fig.suptitle("Wykresy szeregów rozdzielczych erupcji wulkanów")
fig.tight_layout()
fig.subplots_adjust(wspace=0.5)
```

```python

```

```python

```